# Homework 2 (Total Points: 10) <a class="anchor" id="top"></a>


## Submission instructions:

- The assignment is designed for teams of two students.
- Add your solutions to the cells with the following comment `# YOUR CODE HERE`, which denotes that these sections are graded.
- The notebook that you will submit has to have the student ids, separated by underscores (E.g., `12341234_12341234.ipynb`). 
- This will be parsed by a regexp, so please double check your filename.
- Only one member of each group should submit the notebook file to canvas.
- Make sure to check that your notebook runs before submission. A quick way to do this is to restart the kernel and run all the cells.  
- Please do not delete cells. Removing cells can lead to grade deduction. You are allowed to add cells for performing tests or adding auxiliary methods in the given cells.
- Note, that you are not allowed to submit notebooks created with Google Colab. (The reason is that such notebooks do not work with our semi-automated grading system. You may work with Google Colab on developing your solutions. Please then just copy your answers into a new notebook that you haven't opened with Google Colab.)
- Add your name and e-mail in the following cell:

```python
## Add your information in here.
NAME = "Irene Papadopoulou"
NAME2 = "Luuk Verheijen"
EMAIL = "irene.papadopoulou@student.uva.nl"
EMAIL2 = "luuk.verheijen@student.uva.nl"
```

**Assignment instructions Goals**:

In this assignment, you will program several algorithms related to BombKnightSudoku (a variant of Sudoku) on a $(k*k)^2$ size grid, for varying sizes of $k$. You will use several problem solving and search methods.

The homework assignment consists of several parts (a)–(d). You get 1 point for free. For the rest parts you need to develop the following algorithms for solving sudoku game:

 - Part (a) 4pts: solving puzzles using SAT
 - Part (b) 4pts: solving puzzles using ASP
 - Part (c) 2pts: generating puzzles with unique solutions using ASP (or you may opt to solve an easier variant for a maximum of 1 point)

**Game information: BombKnightSudoku's**

Here follows a description of inputs and solutions for **BombKnightSudoku's**. These notions are similar (but not exactly the same) as for the Sudoku's that you might have seen elsewhere, generalized to arbitrary values of k (typically, $k=3$).

An input consists of a $k$-by-$k$ grid of blocks, where each consisting of a $k$-by-$k$ grid of cells, forming a $(k*k)$-by-$(k*k)$ grid of cells. Each cell can (but must not) contain a value between 1 and $k*k$.

An example is the following grid, for $k$=3, written out as $k*k$ lines, each consisting of $k*k$ numbers, separated by spaces. Here the value $0$ represents that a cell does not contain a value between $1$ and $k*k$.

```
0 0 0 0 0 0 0 0 0
0 3 0 0 8 0 0 0 0
0 0 0 0 0 0 0 0 0
0 0 0 0 4 0 0 0 5
0 0 0 0 0 0 0 2 7
0 0 0 0 0 0 3 0 0
0 0 0 0 0 0 9 0 0
0 0 0 0 5 6 0 0 0
0 0 0 0 0 0 0 0 0
```

Here is the same input, where now the $k$-by-$k$ blocks are separated from each other with whitespace.

```
0 0 0  0 0 0  0 0 0
0 3 0  0 8 0  0 0 0
0 0 0  0 0 0  0 0 0

0 0 0  0 4 0  0 0 5
0 0 0  0 0 0  0 2 7
0 0 0  0 0 0  3 0 0

0 0 0  0 0 0  9 0 0
0 0 0  0 5 6  0 0 0
0 0 0  0 0 0  0 0 0
```

A solution for a sudoku input is a $(k*k)$-by-$(k*k)$ grid (for the same $k$), where:

 - Each cell contains a value between 1 and $k*k$.
If a cell ($i$,$j$) contains a value u in the input, then the cell ($i$,$j$) in the solution must contain the same value u.
 - Each two different cells in the same row must contain different values.
 - Each two different cells in the same column must contain different values.
 - Each two different cells in the same $k*k$ block must contain different values.
 - Each two different cells that are directly adjacent (horizontally, diagonally or vertically) must contain different values. In other words, a cell cannot contain the same value as any of the eight cells surrounding it.
 - Each two different cells that can be reached from each other with a single knight's move (as in Chess) must contain different values.

(The last two constraints are what distinguish BombKnightSudoku's from regular Sudoku's.)

For example, a solution for the input mentioned above is the following:

```
5 6 2 7 3 4 1 8 9
7 3 4 1 8 9 5 6 2
1 8 9 5 6 2 7 3 4
6 2 7 3 4 1 8 9 5
3 4 1 8 9 5 6 2 7
8 9 5 6 2 7 3 4 1
2 7 3 4 1 8 9 5 6
4 1 8 9 5 6 2 7 3
9 5 6 2 7 3 4 1 8
```

**Representation of inputs and outputs**

Inputs are represented as a list of lists of numbers.
For example, the BombKnightSudoku input above is represented as:

```
input =
    [[0, 0, 0, 0, 0, 0, 0, 0, 0],
     [0, 3, 0, 0, 8, 0, 0, 0, 0],
     [0, 0, 0, 0, 0, 0, 0, 0, 0],
     [0, 0, 0, 0, 4, 0, 0, 0, 5],
     [0, 0, 0, 0, 0, 0, 0, 2, 7],
     [0, 0, 0, 0, 0, 0, 3, 0, 0],
     [0, 0, 0, 0, 0, 0, 9, 0, 0],
     [0, 0, 0, 0, 5, 6, 0, 0, 0],
     [0, 0, 0, 0, 0, 0, 0, 0, 0]]
```

Solutions are represented similarly, for example:

```
solution =
    [[5, 6, 2, 7, 3, 4, 1, 8, 9],
     [7, 3, 4, 1, 8, 9, 5, 6, 2],
     [1, 8, 9, 5, 6, 2, 7, 3, 4],
     [6, 2, 7, 3, 4, 1, 8, 9, 5],
     [3, 4, 1, 8, 9, 5, 6, 2, 7],
     [8, 9, 5, 6, 2, 7, 3, 4, 1],
     [2, 7, 3, 4, 1, 8, 9, 5, 6],
     [4, 1, 8, 9, 5, 6, 2, 7, 3],
     [9, 5, 6, 2, 7, 3, 4, 1, 8]]
```

Note that some test_inputs are given in the auxiliary python file `sudoku_func.py` which are used for testing your methods.

The necessary libraries for developing all necessary code can be found in the following block:

```python
# Necessary libraries that will be employed
import itertools
from copy import deepcopy
from dataclasses import dataclass
from typing import List

from sudoku_func import *

#!{sys.executable} -m pip install clingo
#!{sys.executable} -m pip install python-sat
import clingo
import pysat
from pysat.formula import CNF
from pysat.solvers import MinisatGH
```

**N.B.** There are two *pysat* libraries. Please make sure to pip install python-sat (*not* pysat).

**N.B.** Please make sure *not* to remove or change the code you have been given (e.g. do not add, modify or remove the function names, names of the parameters, etc.).

## (a): solving puzzles using SAT

Again, implement an algorithm that takes as input a BombKnightSudoku input (for arbitrary `k`), and that outputs a generator that iterates over **all** solutions for this input.

For this assignment, your algorithm should use the approach of encoding the BombKnightSudoku input into a propositional CNF formula &varphi;, calling a SAT solver to find satisfying assignments for &varphi;, and translating these assignments to solutions for the BombKnightSudoku input.

<!-- grade_id: cell-940a32cfcc687489, points: 0, grade: false, solution: true -->
```python
import numpy as np
from pprint import pprint

def range_incl(start, stop):
    return range(start, stop+1)

def pretty_repr_clauses(clauses, p_inv):
    pretty_clauses = []
    for i, clause in enumerate(clauses):
        pretty_clauses.append([])
        for l in clause:
            neg = "~" if np.sign(l) == -1 else ""
            
            pos = p_inv[np.abs(l)]
            
            pretty_clauses[i].append(f"{neg}p_{pos[0]},{pos[1]},{pos[2]}")
    return pretty_clauses

def assignment_to_sudoku(assignment, k, p_inv):
    sudoku = np.zeros((k**2, k**2), dtype=int)
    for l in assignment:
        negated = True if np.sign(l) == -1 else False
        if negated:
            continue
            
        pos = p_inv[np.abs(l)]
        sudoku[pos[0]-1][pos[1]-1] = pos[2]
        
    return sudoku

def create_vars_CNF(k):
    counter = 1
    p = {}
    p_inv = {}
    for i in range_incl(1, k**2):
        p[i] = {}
        for j in range_incl(1, k**2):
            p[i][j] = {}
            for l in range_incl(1, k**2):
                p[i][j][l] = counter
                p_inv[counter] = (i,j,l)
                counter += 1
                
    return p, p_inv

def create_rules_CNF(sudoku, k, p):
    clauses = []
    for r, row in enumerate(sudoku):
        for c, value in enumerate(row):
            
            # Force every cell between 1 and k^2
            clause = []
            for l in range_incl(1, k**2):
                clause.append( p[r+1][c+1][l] )
            clauses.append(clause)
            
            # Keep input
            if value == 0:
                continue
            clauses.append( [ p[r+1][c+1][value] ] )
    
    # Row rule
    for i in range_incl(1, k**2):
        for j in range_incl(1, k**2):
            for l in range_incl(1, k**2):
                for c in range_incl(j+1, k**2):
                    clauses.append( [ -p[i][j][l], -p[i][c][l] ] )
                    
    # Column rule
    for i in range_incl(1, k**2):
        for j in range_incl(1, k**2):
            for l in range_incl(1, k**2):
                for r in range_incl(i+1, k**2):
                    clauses.append( [ -p[i][j][l], -p[r][j][l]] )
                    
    # Block rule
    for br in range_incl(1, k):
        r_start = 1+(br-1)*k
        for bc in range_incl(1, k):
            c_start = 1+(bc-1)*k
            for l in range_incl(1, k**2):
                for n in range_incl(0, k**2-1):
                    i = r_start + n //k
                    j = c_start + n % k
                    
                    for m in range_incl(n+1, k**2-1):
                        r = r_start + m //k
                        c = c_start + m % k

                        clauses.append( [ -p[i][j][l], -p[r][c][l] ] )
    
    # Bomb and Knight rules
    bomb_rel_pos = [(i,j) for i in [-1,0,1] for j in [-1,0,1] if (i,j) != (0,0)]
    knight_rel_pos = [(2,1),(2,-1),(-2,1),(-2,-1),(1,2),(1,-2),(-1,2),(-1,-2)]
    
    for i in range_incl(1, k**2):
        for j in range_incl(1, k**2):
            for l in range_incl(1, k**2):
                for rel_pos in bomb_rel_pos + knight_rel_pos:
                    r = i + rel_pos[0]
                    c = j + rel_pos[1]
                    if r > k**2 or r < 1 or c > k**2 or c < 1:
                        continue

                    clauses.append( [ -p[i][j][l], -p[r][c][l] ] )
                    
    return clauses

def get_single_assignment(clauses):
    solver = MinisatGH(bootstrap_with=clauses)
    solution = solver.solve()
    # print(f"Solvable: {solution}")
    assignment = solver.get_model()
    
    return assignment


def solve_sudoku_SAT(sudoku, k):
    """
    Solve a sudoku input by encoding the problem into SAT, calling a SAT solver, and retrieving the solution for the sudoku input from a satisfying truth assignment given by the SAT solver.

    Parameters:
        sudoku (list(list(int))): An input puzzle.
        k (int): The dimension of the sudoku input.

    Returns:
        A generator using yield instead of returning a (list(list(list(int)))).
    """

    # YOUR CODE HERE
    p, p_inv = create_vars_CNF(k)
    clauses = create_rules_CNF(sudoku, k, p)
    # pprint(pretty_repr_clauses(clauses, p_inv))
    
    while True:
        assignment = get_single_assignment(clauses)
        
        if not assignment:
            break
        
        clauses.append( [-l for l in assignment] )
        
        yield assignment_to_sudoku(assignment, k, p_inv)
```

```python
for i, inp in enumerate(test_inputs):
    if i >= 3:
        break

    print("\nINPUT:")
    print(pretty_repr(inp.sudoku, inp.k))    

    for solution in solve_sudoku_SAT(inp.sudoku, inp.k):
        print("SOLUTION:")
        print(pretty_repr(solution, inp.k))
        print("VALID: ", check_solution(inp.sudoku, inp.k, solution))
```

The following code will test whether your provided solution has the correct number of solution for the test puzzles.

<!-- grade_id: cell-4f0c31f4bfc245aa, points: 4, grade: true, solution: false -->
```python
# This cell will test whether your code for solve_sudoku_SAT()
# finds the right number of solutions on several test inputs.

from sudoku_func import check_num_solutions
from sudoku_func import test_inputs

## Run tests
for num, test_input in enumerate(test_inputs, 1):
    print("- Testing on test input #{}".format(num))
    assert check_num_solutions(
        test_input.sudoku,
        test_input.k,
        test_input.num_solutions,
        solve_sudoku_SAT
    ) == True
    print("  Found exactly the set of correct solutions")

# (There will also be some hidden tests..)

```

## (b): solving puzzles using ASP

Yet again, implement an algorithm that takes as input a BombKnightSudoku input (for arbitrary `k`), and that outputs a generator that iterates over **all** solutions for this input.

For this assignment, your algorithm should use the approach of encoding the BombKnightSudoku input into an answer set program *P*, calling an ASP solver (i.e., clingo) to find answer sets for *P*, and translating these answer sets to solutions for the BombKnightSudoku input.

<!-- grade_id: cell-e08fbc765eba6e6d, points: 0, grade: false, solution: true -->
```python
def extract_answers(program, num_solutions=0):
    
    # Load the answer set program, and call the grounder
    control = clingo.Control(arguments=["--project"])
    control.add("base", [], program)
    control.ground([("base", [])])
    
    # This function transforms the answer set into a list of pairs
    def answer_set_to_list(model):
        output = []
        for atom in model.symbols(shown=True):
            output.append(tuple([arg.number for arg in atom.arguments]))
        return output
                
    # Ask clingo to find all models (using an upper bound of 0 gives all models)
    control.configuration.solve.models = num_solutions
    
    # Call the clingo solver, transform each answer
    for answer in control.solve(yield_=True):
        yield answer_set_to_list(answer)

def answer_set_to_sudoku(answer_set, k):
    sudoku = np.zeros((k**2, k**2), dtype=int)
    for i, j, l in answer_set:            
        sudoku[i-1][j-1] = l
        
    return sudoku


def create_vars_ASP(k):
    return f"""\n% Define vars
    #const k = {k}.
    i(1..k*k).
    j(1..k*k).
    num(1..k*k).
    b(1..k).
    """

def create_rules_ASP(sudoku):
    # Choose exactly one
    rules = """\n% Position rules
    p(I,J,L) : num(L) :- i(I), j(J). % Choose exactly one number per position
    """
    
    # Keep input
    rules += "\n% Keep input\n"
    for r, row in enumerate(sudoku):
        for c, value in enumerate(row):
            if value == 0:
                continue   
            rules += f"    :- not p({r+1},{c+1},{value}).\n"
    
    # Rows and columns rule
    rules += """\n% Restrict rows and columns
    :- p(I,J,L), p(I,C,L), J != C. % Restrict rows
    :- p(I,J,L), p(R,J,L), I != R. % Restrict columns
    """
    
    # Block rule
    rules += """\n% Block rule
    block(I,J,BR,BC) :- b(BR), b(BC), 
                        (BR-1)*k < I, I <= k*BR, i(I),
                        (BC-1)*k < J, J <= k*BC, j(J).
    :- p(I,J,L), p(R,C,L), block(I,J,BR,BC), block(R,C,BR,BC), R != I, C != J.
    """
    
    # Bomb and Knight rules
    bomb_rel_pos = [(i,j) for i in [-1,0,1] for j in [-1,0,1] if (i,j) != (0,0)]
    knight_rel_pos = [(2,1),(2,-1),(-2,1),(-2,-1),(1,2),(1,-2),(-1,2),(-1,-2)]
    
    rules += "\n% Bomb constraints\n"    
    for rel_pos in bomb_rel_pos:
        i_add = rel_pos[0]
        j_add = rel_pos[1]
        rules += f"    :- p(I,J,L), p(R,C,L), R = I+{i_add}, C = J+{j_add}.\n"
        
    rules += "\n% Knight constraints\n" 
    for rel_pos in knight_rel_pos:
        i_add = rel_pos[0]
        j_add = rel_pos[1]
        rules += f"    :- p(I,J,L), p(R,C,L), R = I+{i_add}, C = J+{j_add}.\n"        
    
    return rules

def solve_sudoku_ASP(sudoku, k):
    """
    Solve a sudoku input by encoding the problem into ASP, calling an ASP solver, and retrieving the solution for the sudoku input from an answer set given by the ASP solver.

    Parameters:
        sudoku (list(list(int))): An input puzzle.
        k (int): The dimension of the sudoku input.

    Returns:
        A generator using yield instead of returning a (list(list(list(int)))).
    """

    # YOUR CODE HERE
    program = "#show p/3.\n"
    program += create_vars_ASP(k)
    program += create_rules_ASP(sudoku)
    # print(program)
    
    for answer_set in extract_answers(program):
        yield answer_set_to_sudoku(answer_set, k)
```

```python
for inp in test_inputs:
    # print("INPUT:")
    # print(pretty_repr(inp.sudoku, inp.k))    
    for solution in solve_sudoku_ASP(inp.sudoku, inp.k):    
        # print("SOLUTION:")
        # print(pretty_repr(solution, inp.k))
        print("VALID: ", check_solution(inp.sudoku, inp.k, solution))
```

The following code will test whether your provided solution has the correct number of solution for the test puzzles.

<!-- grade_id: cell-d58b222e808c0f0e, points: 4, grade: true, solution: false -->
```python
# This cell will test whether your code for solve_sudoku_ASP()
# finds the right number of solutions on several test inputs.

from sudoku_func import check_num_solutions
from sudoku_func import test_inputs

## Run tests
for num, test_input in enumerate(test_inputs, 1):
    print("- Testing on test input #{}".format(num))
    assert check_num_solutions(
        test_input.sudoku,
        test_input.k,
        test_input.num_solutions,
        solve_sudoku_ASP
    ) == True
    print("  Found exactly the set of correct solutions")

# (There will also be some hidden tests..)

```

## (c): generating puzzles with unique solutions using ASP

Implement an algorithm that generates `9`-by-`9` BombKnightSudoku inputs (i.e., `k=3`) that:
1. have exactly 10 cells filled with a number between `1` and `9`,
1. and have a unique solution.

That is, your algorithm should output a generator that iterates over BombKnightSudoku inputs with the required properties.

Your algorithm should use the approach of translating this problem to an answer set program *P*, calling an ASP solver (i.e., clingo) to find answer sets for *P*, and translating these answer sets to BombKnightSudoku inputs with the required properties.

(To be clear: you may not 'hardcode' a list of BombKnightSudoku inputs in the answer set program *P*.)

There are two variants of this assignment: a harder variant (d.I) and an easier variant (d.II). You may choose which variant you would like to solve.
If you opt for the harder variant (d.I), you can get a maximum of 2 points for this assignment.
If you opt for the easier variant (d.II), you can get a maximum of 1 points for this assignment.


### (c.I) Harder variant (2pts):

To solve the harder variant, your answer set program *P* should be such that for **each** BombKnightSudoku input *I* with the required properties, there is an answer set of *P* that corresponds to *I*. For this, you should express in *P* the constraints 1. (that exactly 10 cells are filled) and 2. (that the input has a unique solution).

*Hint:* Use the technique of saturation to express constraint 2.

### (c.II) Easier variant (1pt):

To solve the easier variant, your answer set program *P* only needs to be such that for **a subset of** BombKnightSudoku inputs *I* with the required properties, there is an answer set of *P* that corresponds to *I*. For this, you could express in *P*, in addition to the constraint 1. (that exactly 10 cells are filled), some other property that entails constraint 2. (that the input has a unique solution).

*Hint:* If a Sudoku input can be solved using some [basic strategies](http://pi.math.cornell.edu/~mec/Summer2009/Mahmood/Solve.html), it must have a unique solution. Consider variants of these basic strategies for BombKnightSudoku.

<!-- grade_id: cell-fb57584e65e58ba1, points: 0, grade: false, solution: true -->
```python
def create_rules_incorrect():
    rules = ""

    # Keep input
    rules += """\n% Keep input
    incorrect :- sudoku(I,J,L1), other_solution(I,J,L2), L1 != L2.
    """
    
    # Rows and columns rule
    rules += """\n% Restrict rows and columns
    incorrect :- other_solution(I,J,L), other_solution(I,C,L), J != C. % Restrict rows
    incorrect :- other_solution(I,J,L), other_solution(R,J,L), I != R. % Restrict columns
    """
    
    # Block rule
    rules += """\n% Block rule
    incorrect :- other_solution(I,J,L), other_solution(R,C,L), block(I,J,BR,BC), block(R,C,BR,BC), R != I, C != J.
    """
    
    # Bomb and Knight rules
    bomb_rel_pos = [(i,j) for i in [-1,0,1] for j in [-1,0,1] if (i,j) != (0,0)]
    knight_rel_pos = [(2,1),(2,-1),(-2,1),(-2,-1),(1,2),(1,-2),(-1,2),(-1,-2)]
    
    rules += "\n% Bomb constraints\n"    
    for rel_pos in bomb_rel_pos:
        i_add = rel_pos[0]
        j_add = rel_pos[1]
        rules += f"    incorrect :- other_solution(I,J,L), other_solution(R,C,L), R = I+{i_add}, C = J+{j_add}.\n"
        
    rules += "\n% Knight constraints\n" 
    for rel_pos in knight_rel_pos:
        i_add = rel_pos[0]
        j_add = rel_pos[1]
        rules += f"    incorrect :- other_solution(I,J,L), other_solution(R,C,L), R = I+{i_add}, C = J+{j_add}.\n"  

    # Make sure it's a new solution
    rules += """\n% other_solution cannot be the same as the initial solution
    incorrect :- other_solution(I,J,L) : p(I,J,L).
    """    

    return rules


def generate_puzzles_with_unique_solutions(k=3, show_n=10, sudoku=None):
    """
    Algorithm that generates BombKnightSudoku puzzles with a unique solution
    and where exactly 10 cells are filled.
    
    Returns:
        A generator using yield instead of returning a (list(list(list(int)))).
    """

    # YOUR CODE HERE
    if not sudoku:
        sudoku = np.zeros((k**2, k**2))

    # program = "#show sudoku/3.\n"
    program = "#show sudoku/3."
    program += create_vars_ASP(k)
    program += create_rules_ASP(sudoku)

    # Show only n cells
    program += f"""\n% Show only n cells
    {show_n} {{ show(I,J) : i(I), j(J) }} {show_n}.
    sudoku(I,J,L) :- p(I,J,L), show(I,J).
    """

    # # Show only n cells by hiding (k**2)**2-n cells
    # hide_n = (k**2)**2-show_n
    # program += f"""\n% Show only n cells
    # {hide_n} {{ hide(I,J) : i(I), j(J) }} {hide_n}.
    # sudoku(I,J,L) :- p(I,J,L), not hide(I,J).
    # """

    # # Show only n cells by generating all hidden subsets and then checking
    # # that n are shown.
    # program += f"""\n% Show only n cells
    # {{ hide(I,J) : i(I), j(J) }}.
    # sudoku(I,J,L) :- p(I,J,L), not hide(I,J).

    # :- not {show_n} {{ sudoku(I,J,L) : i(I), j(J), num(L) }} {show_n}.
    # """

    # Generate the search space for other solutions
    program += """\n% Generate the search space for other solutions
    other_solution(I,J,L) : num(L) :- i(I), j(J).
    """

    # State which solutions are incorrect
    program += create_rules_incorrect()

    # Finish our saturation
    program += """\n% Saturate
    saturate :- incorrect.
    other_solution(I,J,L) :- i(I), j(J), num(L), saturate. % if incorrect we saturate
    :- not saturate.
    """

    # print(program)

    for answer_set in extract_answers(program):
        yield answer_set_to_sudoku(answer_set, k)

    
```

```python
k = 3

for i, puzzle in enumerate(generate_puzzles_with_unique_solutions(k=k)):
    if i >= 10:
        break

    print(f"Puzzle {i+1}:\t", check_num_solutions(puzzle, k, 1, solve_sudoku_ASP))
    # print(pretty_repr(puzzle, k))
```

<!-- grade_id: cell-d86460df48d77382, points: 2, grade: true, solution: false -->
```python
# This cell will test whether your function yields
# (at least) 10 different sudoku puzzles with unique solutions.

# (These tests are hidden.)

```